# Структуры данных

Как только мы начнем работать со множеством однотипных данных, нам будет удобней хранить данные в структурах, таких как массивы или словари (а не просто создавать множество переменных). 

Типы охватываемых структур данных : 
- [Кортежи](#Кортежи)
- [Именованные кортежи](#Именованные-кортежи)
- [Словари](#Словари)


## Кортежи (Tuples)

Мы можем создать кортеж, заключив упорядоченную коллекцию элементов в `( )`.

Синтаксис: 

```julia
(item1, item2, ...)
```

```julia
favoritelang = ("Python","Julia","R")
```

```julia
typeof(favoritelang)
```

```julia
myfavoriteanimals = ("penguins", "cats", "sugargliders")
```

### Indexing
Мы можем индексировать в этот кортеж,

```julia
myfavoriteanimals[1]
```

### Immutability
Кортежи неизменны, мы не можем обновить их

```julia
myfavoriteanimals[1] = "otters"
```

### Iteration
Мы можем перебирать кортеж.

```julia
for lang in favoritelang
    print(lang)
    print(", ")
end
```

### Несколько полезных методов (общих для всех структур данных)
- `isempty` - проверяет, пуста ли структура данных,
- `length` - возвращает длину структуры данных,
- `in` - проверяет принадлежность к элементу,
- `unique` - возвращает коллекцию уникальных элементов,
- `reduce` - свёртывает данную коллекцию` itr` с заданным бинарным оператором `op`,
- `maximum` (или` minimum`) - возвращает наибольший (или наименьший) результат вызова функции `fun` для каждого элемента` itr`

```julia
isempty(favoritelang)
```

```julia
length(favoritelang)
```

```julia
"Java" in favoritelang
```

```julia
"Julia" in favoritelang
```

```julia
favoritelang2 = ("Python","Julia","R","Julia","Python");
println("Printing all the elements:")
for lang in favoritelang2
    print(lang);print(", ")
end
println("\n\nNow printing only the unique elements:")
for lang in unique(favoritelang2)
    print(lang);print(", ")
end
```

---
Свёртка кортежа (см. [Свёртка списка](https://ru.wikipedia.org/wiki/%D0%A1%D0%B2%D1%91%D1%80%D1%82%D0%BA%D0%B0_%D1%81%D0%BF%D0%B8%D1%81%D0%BA%D0%B0))

```julia
reduce(*,favoritelang)
```

```julia
t = (2,3,-4);
reduce(+,t)
```

```julia
# Python is the longest string, so we expect the answer to be 6
maximum(length,favoritelang)
```

```julia
# R is the shortest string, so we expect the answer to be 1
minimum(length,favoritelang)
```

## Именованные кортежи

Как вы могли догадаться, `NamedTuple`ы это такие же `Tuple`ы за исключением того, что каждый элемент дополнительно имеет имя! У них есть специальный синтаксис, использующий `=` внутри кортежа:

```julia
(name1 = item1, name2 = item2, ...)
```

```julia
myfavoriteanimals = (bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

```julia
typeof(myfavoriteanimals)
```

Как и обычные `Tuples`, `NamedTuples` упорядочены, так что мы можем получить их элементы с помощью индексации:

```julia
myfavoriteanimals[1]
```

Они также добавляют особую возможность доступа к значениям по их имени:

```julia
myfavoriteanimals.bird
```

## Словари

Если у нас есть наборы данных, связанных друг с другом, мы можем сохранить эти данные в словаре. Мы можем создать словарь, используя функцию `Dict()`.

Syntax:

```julia
Dict(key1 => value1, key2 => value2, ...)
```

Хорошим примером является список контактов, где мы связываем имена с номерами телефонов.

```julia
phonebook = Dict("Jenny" => "867-5309", "Ghostbusters" => "555-2368")
```

В этом примере каждое имя и номер являются парой «ключ» и «значение». Мы можем получить номер Дженни (значение), используя связанный ключ

```julia
phonebook["Jenny"]
```

Мы можем добавить еще одну запись в этот словарь следующим образом

```julia
phonebook["Kramer"] = "555-FILK"
```

Давайте посмотрим, как сейчас выглядит наша телефонная книга...

```julia
phonebook
```

Мы можем удалить Крамера из нашего списка контактов и одновременно получить его номер, используя `pop!`

```julia
pop!(phonebook, "Kramer")
```

```julia
phonebook
```

В отличие от кортежей и массивов, словари не упорядочены. Таким образом, мы не можем индексировать их.

```julia
phonebook[1]
```

В этом примере `julia` думает, что вы пытаетесь получить доступ к значению, связанному с ключом `1`.

### Некоторые полезные методы
- `keys`
- `values`
- `pairs`
- `haskey`

```julia
keys(phonebook)
```

```julia
values(phonebook)
```

```julia
pairs(phonebook)
```

```julia
haskey(phonebook,"Kramer")
```

```julia
haskey(phonebook,"Ghostbusters")
```

### Функция `merge()`
Эта функция создает объединенную коллекцию из заданных коллекций. При необходимости типы результирующей коллекции будут приведены, чтобы соответствовать типам объединенных коллекций. Если такой же ключ присутствует в другой коллекции, значением для этого ключа будет значение, которое он имеет в последней перечисленной коллекции.

```julia
a = Dict("foo" => 0.0, "bar" => 42.0);
b = Dict("baz" => 17, "bar" => 13.0);
```

```julia
a
```

```julia
b
```

```julia
# Note the order of merging
merge(a, b)
```

```julia
# Now the order is reversed
merge(b,a)
```

Если `merge ()` используется с аргументом функции "объединитель", то значения с одним и тем же ключом будут объединены с помощью функции объединения.

```julia
a = Dict("foo" => 0.0, "bar" => 42.0);
b = Dict("baz" => 17, "bar" => 13.0);

# Using a +(sum) function as the combiner
merge(+,a,b)
```

### Упражнения

#### 1

Попробуйте добавить "Emergency" в `phonebook` задав значение `911`.

```julia
phonebook["Emergency"] = 911
```

Что тут не так?

```julia

```

Попробуйте функцию `string(911)`

```julia

```

#### 2
Создайте новый словарь `flexible_phonebook` где номер Дженни будет целочисленным, а у Охотников за привидениями - строковым используя следующий код

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

<!-- grade_id: cell-66f1d636c676c2f2, points: 1, grade: true, solution: false -->
```julia
@assert flexible_phonebook == Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

#### 3
Добавьте туда ключ "Emergency" со значением `911` (целочисленное).

```julia
flexible_phonebook["Emergency"] = 911
```

<!-- grade_id: cell-ee862ab938680259, points: 1, grade: true, solution: false -->
```julia
@assert haskey(flexible_phonebook, "Emergency")
```

<!-- grade_id: cell-c8c40ababe7cdfce, points: 2, grade: true, solution: false -->
```julia
@assert flexible_phonebook["Emergency"] == 911
```

#### 4
Почему в `flexible_phonebook` целочисленное добавляется, а в `phonebook` нет? Как нам следует инициализировать `phonebook`, чтоб он мог принимать целочисленные значения?

```julia

```