# Инструменты - Использование массивов для хранения данных

### Введение в массивы

**Массивы - это наборы блоков, которые используются для хранения данных.** В последнем блокноте мы увидели изображение, которое может помочь нам изобразить одномерный массив:

<img src="images/array_cartoon.png" alt="Drawing" style="width: 500px;"/>

*Почему мы хотим, чтобы подобный объект хранил наши данные?*

Альтернативой использованию массива в некоторых контекстах может быть присвоение каждому отдельному фрагменту данных следующих имен:

```julia
a = 1.1
b = 2.2
c = 3.3
```

Мы можем визуализировать, как хранятся эти данные:

<img src="images/without_arrays.png" alt="Drawing" style="width: 500px;"/>



Это похоже на наличие отдельного блока для каждой части данных, а не серии связанных блоков для всех наших данных. 

Чем больше у нас данных, тем больше раздражает отслеживание всех этих ящиков и их названий. Кроме того, если мы хотим сделать то же самое со многими частями данных, гораздо проще поместить все эти части данных в одно место для одновременной работы с ними. 

Например, мы можем захотеть умножить `a`,` b` и `c` на` 2`. Мы могли бы умножить три раза:

```julia
a * 2
b * 2
c * 2
```

Или вместо этого мы могли бы создать один массив (назовем его `numbers`) и умножить этот массив на` 2`:
```julia
numbers * 2
```

Синтаксис для создания этого массива `numbers`:

```julia
numbers = [a, b, c]
```

Или мы могли бы просто написать

```julia
numbers = [1.1, 2.2, 3.3]
```

Стоит отметить, что в Julia одномерные массивы также называют «векторами».

---
Как и в других языках программирования, под массивами в Julia понимается коллекция упорядоченных элементов, размещённая в многомерной сетке. Векторы и матрицы являются частными случаями массивов.


В Julia предусмотрено достаточно много способов создать массив (с учётом типа, значения, размерности и пр.), но по моему скромному мнению в большинстве случаев удобнее пойти путём Matlab-подобных сред и инициализировать массив через перечисление его элементов в квадратных скобках, разделяя строки знаком " ; ". Чтобы обратиться к элементу, укажите его индекс в квадратных скобках, если массив многомерный, перечислите индексы через запятую. ***Индексация начинается с единицы.*** Для облегчения доступа индекс последнего элемента хранится в переменной `end`. 

### Создание массивов

Мы можем создать массив `numbers`, собрав наши элементы,` a`, `b` и` c` (или `1.1`,` 2.2` и `3.3`) в квадратные скобки.

Синтаксис:

```julia
[item1, item2, ...]
```

Создание пустого массива с абстрактным типом

```julia
my_array = []
```

Создание пустого массива с конкретным типом

```julia
my_array = (Integer)[]
```

```julia
my_array = (Float64)[]
```

Массив из трех элементов (здесь тип будет определен автоматически Julia)

```julia
array1 = [1, 2, 3]
```

Это вектор-столбец. Мы также можем создать вектор-строку, опуская запятые.

```julia
array2 = [1 2 3]
```

Если мы смешиваем различные типы элементов в массиве, тогда массивы примут тип элемента, который является *более общим*. В приведенном ниже примере два значения будут иметь тип `Int64`, а последнее - тип `Float64`. Массив вернет тип `Float64`.

```julia
array3 = [1, 2, 3.0]
```

```julia
array4 = ["one",pi,42];
for i in array4
    println("Type of \'$i\' is $(typeof(i))")
end
```

### Многомерные массивы

Мы можем создавать массивы, состоящие из более чем одной строки или столбца элементов. Мы можем создавать многомерные массивы. Для этого вставляем квадратные скобки и манипулируем запятыми и точками с запятой.

Обратите внимание на отсутствие запятых между внутренними вложенными квадратными скобками

```julia
array5 = [[1, 2, 3] [4, 5, 6] [7, 8, 9]]
```

Если мы хотим сначала заполнить элементы по строкам, мы используем точку с запятой

```julia
array6 = [[1 2 3]; [4 5 6]; [7 8 9]]
```

#### Упражнение



Создайте массив с именем `first_array`, в котором будут храниться числа от 10 до 20. 

*Напоминание: ESC + b, чтобы открыть поле ниже.*

Одним из способов создания многомерного массива является использование функции `rand`. `rand` принимает произвольное количество аргументов, где количество аргументов определяет количество измерений, которые будет иметь массив! 

Например, если мы передадим `rand` два целочисленных ввода, он сгенерирует двумерный массив, а три целочисленных ввода сгенерируют трехмерный массив:

```julia
rand(4, 3)
```

```julia
rand(4, 3, 2)
```

### Массивы массивов

Массивы также могут хранить другие массивы.
<br><br>Например, ниже приведены массивы массивов:


```julia
favorites = [ ["koobideh", "chocolate", "eggs"], ["penguins", "cats", "sugargliders"] ]
```

```julia
numbers = [ [1, 2, 3], [4, 5], [6, 7, 8, 9] ]
```

### Создание массивов включениями (comprehensions)

В качестве альтернативы мы можем создавать массивы с помощью *'включений'*. Это хороший способ автоматизировать создание массива, если вы не хотите вводить длинные списки чисел в квадратных скобках. Во включении вы пишете код в квадратных скобках, который будет генерировать массив при выполнении этого кода. Давайте посмотрим на пример.

```julia
counting = [2 * i for i in 1:10]
```

В приведенной выше строке кода говорится, что мы хотим создать массив с именем `counting`, в котором каждое из целых чисел от 1 до 10 будет храниться умноженным на два. Это включение состоит из нескольких различных частей, поэтому давайте разберем их:

<img src="images/array_comprehension.png" alt="Drawing" style="width: 500px;"/>

Если бы мы захотели создать еще один массив `roots`, в котором хранятся квадратные корни всех целых чисел от 1 до 10, мы могли бы изменить приведенный выше код так, чтобы

```julia
roots = [sqrt(i) for i in 1:10]
```

---
Создать массив с элементами $3 x^2$, где $x$ - нечетное число от 1 до 9 (включительно)

```julia
array15 = [3*i^2 for i in 1:2:9]
```

Можно использовать несколько переменных для инициализации двумерного массива

```julia
array16 = [a+2b for a in -1:1, b in -2:2]
```

Массив квадратов элементов, если квадрат не делится на 5 или 4 (обратите внимание, что используется `&&`)

```julia
array17=[i^2 for i=1:10  if (i^2%5!=0 && i^2%4!=0)]
```

---
#### Упражнение

Создайте массив `squares`, в котором будут храниться квадраты всех целых чисел от 1 до 100.

### Глядя внутрь массива

Мы можем «индексировать» массив, чтобы получить содержимое внутри массива в определенной позиции. Для индексации в нашем массиве `counting` мы помещаем квадратные скобки после имени массива и помещаем номер позиции элемента / данных, которые мы хотим, в эти скобки. Например, мы могли бы сказать,

```julia
counting[3]
```

чтобы получить третий элемент в массиве под названием `counting`.

#### Упражнение 2

Выполните код в следующей ячейке. Он сгенерирует массив myprimes, в котором будут храниться первые 168 простых чисел. Воспользуйтесь индексом в `myprimes`, чтобы получить 89-е наименьшее простое число. Что это за число?

```julia
import Pkg
Pkg.add("Primes")
```

```julia
using Primes
```

```julia
myprimes = primes(1000); # Точка с запятой подавляет вывод, попробуйте удалить его
```

### Базовые функций для работы с массивами

- `length(A)` - число элементов.
- `ndims(A)` - число размерностей.
- `size(A)` - кортеж размерностей.
- `size(A, n)` - размерность в заданном направлении.
- `copy(A)` - создание копии массива. 

```julia
array6
```

Проверяем, сколько у нас элементов

```julia
length(array6)
```

Проверка размера (размерности)

```julia
size(array6)
```

Создание трехмерного массива случайных чисел с плавающей запятой

```julia
a = rand(3,2,2)
```

```julia
ndims(a)
```

### Специальные случаи массивов - `ones`, `zeros`, `fill`

Специальные функции, такие как `one()` и `zeros()`, могут использоваться для создания массивов с единицами и нулями соответственно. Другая функция `fill()` может использоваться для заполнения массива заранее определенным значением.

```julia
ones(5)
```

```julia
ones(2,3)
```

```julia
zeros(4)
```

```julia
zeros(3,3)
```

```julia
fill([1,2],3)
```

```julia
fill(3.5,(3,2))
```

```julia
fill("Julia",(2,3))
```

### Изменение формы, транспонирование, индексирование

```julia
a = collect(1:12)
```

```julia
reshape(a,(2,6))
```

```julia
reshape(a,(3,4))
```

```julia
b = reshape(a,(3,4))
```

```julia
c=transpose(b)
```

```julia
b'
```

---
Мы могли бы транспонировать массив `array1`, что сделало бы его похожим на массив `array2`. Помните, что действие транспонирования меняет местами строки и столбцы.

```julia
transpose(array1)
```

```julia
array1'
```

```julia
array1' == array2
```

Но тип транспонированного массива не `Array`. Поскольку транспонирование происходит из модуля `LinearAlgebra`, тип массива изменяется. Мы можем использовать оператор строгого равенства `===`, чтобы проверить это.

```julia
typeof(array1')
```

```julia
array1'===array2
```

---
Если мы хотим получить значение из двумерного массива, мы индексируем, указывая строку и столбец интересующего элемента. Например, если у нас есть массив `A`, заданный

```julia
A = [ 1  2  3
      4  5  6
      7  8  9 ]
 ```
 мы могли бы взять число 6, сказав
 
 ```
 A[2, 3]
 ```
 
 поскольку 6 находится во 2-м ряду и 3-м столбце `A`.

```julia

```

Индексирование - 2-я строка, 3-й столбец

```julia
c[2,3]
```

Весь 3-й ряд

```julia
c[3,:]
```

Весь 2-ой столбец

```julia
c[:,2]
```

Блок, состоящий из 2-й и 4-й строк и 1-го и 3-го столбца

```julia
c[[2,4],[1,3]]
```

### Срезы

Когда массивы становятся все больше и больше, мы можем захотеть выбрать только его части на основе некоторого правила или правил.

Вместо того, чтобы захватить одно число в массиве, мы также можем взять **срез** массива, то есть подмножество массива, которое может содержать несколько значений. Чтобы взять фрагмент `counting`, который включает 3-ю, 4-ю и 5-ю записи, мы используем следующий синтаксис с двоеточием (`:`):

```julia
counting[3:5]
```

---
Функция `rand()` возвращает случайное значение от 0 до 1. Мы также можем указать значения для выбора. В приведенном ниже примере мы выбираем целые числа в диапазоне \[10, 20] и создаем массив из десяти строк и пяти столбцов.

```julia
array13 = rand(10:20, 10, 5)
```

Выбор всех значений строки в столбце 2

```julia
array13[:, 2]
```

Все значения строк в столбцах 2 и 5

```julia
array13[:, [2, 5]]
```

Все значения строк в столбцах 2, 3 и 4

```julia
array13[:, 2:4]
```

Значения в строках 2, 4, 6 и в столбцах 1 и 5

```julia
array13[[2, 4, 6], [1, 5]]
```

Значения в строке 1 от столбца 3 до последнего столбца

```julia
array13[1, 3:end]
```

---
#### Упражнение

Воспользуйтесь срезом в `myprimes`, чтобы получить наименьшие простые числа от 89-го до 99-го  (включительно).

```julia

```

## Логические операции с массивами
Теперь мы рассмотрим применение некоторых правил. Обратите внимание, как мы используем поэлементный логический оператор в виде точки (обратите внимание на оператор `.>`).

Кроме того, используя функцию `findall ()`, мы можем вернуть значения индекса.

Булева логика (возвращает только `true` и `false`)

```julia
array13[:, 1]
```

```julia
array13[:, 1] .> 14
```

Возврат индекса истинных значений с помощью `findall ()`

```julia
findall(array13[:, 1] .> 14)
```

## Сортировка и min/max

```julia
array13
```

Сортировка по строкам

```julia
sort(array13,dims=1)
```

Сортировка по столбцам

```julia
sort(array13,dims=2)
```

Возвращает максимум (первый, если их больше одного) и его индекс в кортеже

```julia
findmax(array13)
```

```julia
findmin(array13)
```

```julia
argmax(array13)
```

```julia
argmin(array13)
```

## Простые сводные операции
- `sum`
- `prod`
- `mean`
- `std`
- `middle`
- `median`

```julia
array17
```

```julia
sum(array17)
```

```julia
prod(array17)
```

Такие функции, как `mean ()`, `std ()` и `middle ()` были перемещены в модуль `Statistics` стандартной библиотеки; вам может потребоваться сначала ввести `using Statistics` для их использования

```julia
import Pkg
Pkg.add("Statistics")
```

```julia
using Statistics
```

```julia
mean(array17)
```

```julia
std(array17)
```

```julia
middle(array17)
```

```julia
median(array17)
```

Для двумерных массивов мы можем указать размерности, по которым должна выполняться сводная операция.

```julia
array16
```

```julia
std(array16)
```

```julia
mean(array16,dims=1)
```

```julia
mean(array16,dims=2)
```

## Поэлементные операции

```julia
array18 = [2*a+b for a in 1:3, b in -2:2]
```

```julia
array16
```

```julia
array16+array18
```

Обратите внимание на точку перед `*`, чтобы обозначить поэлементную операцию.

```julia
array16.*array18
```

Деление может давать значения `Inf`, если производится деление на ноль

```julia
array18./array16
```

---
Мы даже можем поэлементно возвести один массив в степень другого. Но для того, чтобы это работало, по крайней мере один из них должен быть типа `Float`, а не `Integer`. Итак, мы создаем массив `Float64` и демонстрируем это.

```julia
array19 = (Float64)[a+b for a in 1:3, b in -1:3]
```

```julia
array18.^array19
```

### Модификация массива

В этом разделе мы увидим, как редактировать массив, который мы уже создали. Посмотрим как

1) Обновить элемент в существующей позиции в массиве;

2) Добавить элемент в конец массива;

3) Удалить элемент из конца массива.

#### Обновить элемент в существующей позиции в массиве

Прежде всего, мы можем обновить элемент в существующей позиции в массиве путем индексации в массиве и изменения значения по желаемому индексу. Например, допустим, у нас есть массив с именем `myfriends`:

```julia
myfriends = ["Ted", "Robin", "Barney", "Lily", "Marshall"]
```

Мы можем получить имя Барни из 'myfriends' через

```julia
myfriends[3]
```

и мы можем изменить «Барни» на что-то еще, переназначив `myfriends[3]`:

```julia
myfriends[3] = "Baby Bop"
```

Обратите внимание, что одиночный `=` назначает новую переменную значению слева от знака `=`.

#### Упражнение 4

Используйте включение массива для создания массива `mysequence`, в котором хранятся числа от 4 до 10. Замените в нём последний элемент, `10`, на` 11`.

#### Добавить элемент в конец массива

```julia
counting = [i for i in 1:10]
```

Мы можем добавить элемент в конец массива с помощью функции `push!`. Не беспокойтесь о восклицательном знаке в конце `push!`; мы поговорим о том, почему это там позже, когда будем обсуждать функции. 

А пока обратите внимание, что когда вы вызываете `push!` Для массива и значения, вы добавляете это значение в конец массива. Например,

```julia
push!(counting, 1000)
```

превратит `counting` из 

```julia
[1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
``` 
в 
```julia
[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 1000]
```

Удостоверьтесь, что это работает

```julia

```

#### Упражнение 5

Скопируйте следующий код, чтобы объявить массив `fibonacci`:

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

Используйте `push!`, Чтобы добавить `21` к` fibonacci` после числа `13`.

#### Удалить элемент из конца массива

Чтобы удалить элемент из конца массива, используйте функцию `pop!`. При использовании `pop!` нужен только один входной аргумент, а именно массив, который вы хотите изменить. Например, если

```julia
counting = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 1000]
```

тогда
```julia
pop(counting)
```
удалит `1000` с конца `counting`.

#### Упражнение 6

Используйте `pop!` чтобы удалить `21` из `fibonacci`. Что возвращает этот вызов функции?

```julia

```

#### Упражнение 7

Что является последним элементом `fibonacci`? Попробуйте

```julia
fibonacci[end]
```

```julia

```

#### Упражнение 8
 
 Скопируйте и выполните следующий код, чтобы получить массив myprimematrix, в котором хранятся первые 100 простых чисел.
 
 ```
 myprimematrix = reshape(primes(541), (10, 10))
 ```
 
 Возьмите простое число в 8-й строке и 5-м столбце с помощью индексации. Каково значение этого простого?

```julia

```

### Копирование массивов

Будьте осторожны, когда хотите скопировать массивы!

Допустим, вы хотите скопировать существующий массив, например, «fibonacci», в массив под названием «somenumbers». Помните, как мы определили `fibonacci`:

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

Что если мы попытаемся скопировать Фибоначчи, сказав

```julia
somenumbers = fibonacci
```
?
Выполните этот код, чтобы попытаться скопировать Фибоначчи. Посмотрите на `somenumbers`, чтобы увидеть, хранит ли он то, что вы хотите.

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

```julia
somenumbers = fibonacci
```

```julia
somenumbers
```

Теперь давайте обновим `somenumbers`, изменив его первый элемент.

```julia
somenumbers[1] = 404
```

Теперь давайте посмотрим внутрь Фибоначчи.

```julia
fibonacci
```

#### Упражнение 9

Что является первым элементом в `Фибоначчи`?

### Так как копировать-то?

Скопировали ли мы 'Фибоначчи'?

Нет, к сожалению нет. Когда мы пытались копировать, все, что мы делали, это давали `fibonacci` новое имя, `somenumbers`. Теперь, когда мы обновляем `somenumbers`, мы также обновляем` fibonacci`, потому что они - один и тот же объект!

Если мы действительно хотим сделать *копию* массива, привязанного к `fibonacci`, мы можем использовать функцию `copy`:

```julia
# Во-первых, восстановить Фибоначчи

fibonacci[1] = 1
fibonacci
```

```julia
somemorenumbers = copy(fibonacci)
```

```julia
somemorenumbers[1] = 404
```

```julia
fibonacci
```

В этом последнем примере Фибоначчи не был обновлен. Следовательно, мы видим, что массивы, связанные с `somemorenumbers` и `fibonacci`, различны.

#### Упражнение 10

Скопируйте myprimematrix в mynewprimematrix. Замените `mynewprimematrix [3,3]` на 1234.

```julia

```

### Упражнения

#### 1 
Создайте массив, `a_ray`, Спомощью следующего кода:

```julia
a_ray = [1, 2, 3]
```

Добавьте число `4` в конец масива, а затем удалите его.

```julia

```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3]
```

#### 2 
Попробуйте добавить "Emergency" в `phonebook` задав значение `string(911)`, а потом попробуйте

```julia
phonebook["Emergency"] = 911
```

Что тут не так?

```julia

```

#### 3 
Создайте новый словарь `flexible_phonebook` где номер Дженни будет целочисленным, а у Охотников за привидениями - строковым используя следующий код

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

```julia

```

<!-- grade_id: cell-66f1d636c676c2f2, points: 1, grade: true, solution: false -->
```julia
@assert flexible_phonebook == Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

#### 4
Добавьте туда ключ "Emergency" со значением `911` (целочисленное).

```julia

```

<!-- grade_id: cell-ee862ab938680259, points: 1, grade: true, solution: false -->
```julia
@assert haskey(flexible_phonebook, "Emergency")
```

<!-- grade_id: cell-c8c40ababe7cdfce, points: 2, grade: true, solution: false -->
```julia
@assert flexible_phonebook["Emergency"] == 911
```

#### 5 
Почему в `flexible_phonebook` целочисленное добавляется, а в `phonebook` нет? Как нам следует инициализировать `phonebook`, чтоб он мог принимать целочисленные значения?

```julia

```