# Linear algebra with Julia

---
Создадим массив 4x3 со случайными целыми числами (от 1 до 20)

```julia
a = rand(1:20,(4,3))
```

## Поэлементная сумма

```julia
sum(a)
```

```julia
sum(a,dims=1)
```

```julia
sum(a,dims=2)
```

## Поэлементное произведение

```julia
prod(a)
```

```julia
prod(a,dims=1)
```

```julia
prod(a,dims=2)
```

## Среднее

```julia
import Pkg
Pkg.add("Statistics")
```

```julia
using Statistics
```

```julia
mean(a)
```

```julia
mean(a,dims=1)
```

```julia
mean(a,dims=2)
```

## Транспонирование, диагонализация, след

Вам нужно использовать библиотеку `LinearAlgebra` для большинства функций в этом разделе. Кроме того, для большинства этих функций нам нужна квадратная матрица, поэтому мы создаем её.

```julia
import Pkg
Pkg.add("LinearAlgebra")
```

```julia
using LinearAlgebra
```

```julia
b = rand(1:20,(4,4))
```

```julia
transpose(b)
```

След матрицы, т.е. сумма диагональных элементов

```julia
tr(b)
```

Извлечение диагональных элементов как массива

```julia
diag(b)
```

## Ранг, определитель, инверсия, ядро (nullspace)

```julia
rank(b)
```

```julia
inv(b)
```

```julia
det(b)
```

Для неквадратных матриц, таких как `a`, не существует строгой обратной функции, но существует [псевдообратная функция](https://en.wikipedia.org/wiki/Moore%E2%80%93Penrose_inverse) `pinv`

Псевдообратная функция

```julia
pinv(a)
```

Чтобы получить ядро, нам нужна матрица, которая не является [полноранговой](https://www.cds.caltech.edu/~murray/amwiki/index.php/FAQ:_What_does_it_mean_for_a_non-square_matrix_to_be_full_rank%3F), т.е. мы необходимо иметь линейно зависимые строки/столбцы. Итак, мы создаем такую матрицу `a`

```julia
c = [a 2*a[:,2]]
```

Ранг матрицы `c` меньше ее размерности

```julia
rank(c)
```

Ядро матрицы

```julia
nullspace(c)
```

## Векторные и матричные нормы
[Euclidean norm](https://en.wikipedia.org/wiki/Norm_(mathematics)#Euclidean_norm), ${‖\bf{X}‖}_{2} = \sqrt{x_1^2 + x_2^2 + \ldots + x_n^2}$ ищется с помощью `LinearAlgebra.norm(x)`. 

```julia
X = [2, 4, -5]
```

Евклидова норма

```julia
norm(X)
```

p-норма определяется как
$$
\|A\|_p = \left( \sum_{i=1}^n | a_i | ^p \right)^{1/p}
$$

1-норма (сумма величин элементов, т.е. абсолютных значений)

```julia
norm(X,1)
```

Известно понятие [_infinity norm_](https://en.wikipedia.org/wiki/Norm_(mathematics)#Maximum_norm_(special_case_of:_infinity_norm,_uniform_norm,_or_supremum_norm)) $||\bf{X}||_\infty$, что в основном эквивалентно максимальному значению вектора / матрицы. Конечно, мы не можем достигнуть истинной бесконечности, но если мы возьмем достаточно большую норму, скажем 20, мы увидим, что значение довольно близко к максимальному элементу `X`, т.е. 5. И тогда мы можем использовать` Inf` в качестве аргумента.

```julia
norm(X,20)
```

```julia
norm(X,Inf)
```

Евклидово расстояние между двумя векторами $\bf{X}$ и $\bf{Y}$ задаётся как $||\bf{X}-\bf{Y}||_2$ и может быть легко вычислен в Julia, используя `norm(X-Y)`

```julia
X = [2, 4, -5];
Y = [1,-1,3];
norm(X-Y)
```

Мы можем проверить ответ, используя базовое определение, то есть _квадратный корень из суммы разности элементов_

```julia
sqrt(sum((X-Y).^2))
```

[Угол между двумя векторами](https://onlinemschool.com/math/library/vector/angl/) $\bf{X}$ и $\bf{Y}$ задаётся как: 
$$\text{cos}^{-1}\frac{X^TY}{||X||_2||Y||_2}$$

```julia
acos((transpose(X)*Y)/(norm(X)*norm(Y)))
```

Для двумерных матриц для вычисления так называемой [операторной нормы]() нам нужно использовать функцию `opnorm ()`. $1$-норма матрицы (максимальная абсолютная сумма столбцов):

```julia
d = [5 -4 2 ; -1 2 3; -2 1 0]
```

```julia
opnorm(d,1)
```

Бесконечная норма (максимальная абсолютная сумма строк)

```julia
opnorm(d,Inf)
```

По умолчанию вычисляется евклидова норма 

```julia
opnorm(d)
```

## Масштабирование, вращение, обращение

```julia
d
```

Обратите внимание, что это функция изменяет исходную матрицу
$$
d = d \times 2
$$

```julia
rmul!(d,2)
```

```julia
d
```

```julia
d = [5 -4 2 ; -1 2 3; -2 1 0]
```

Поворот на 180 градусов

```julia
rot180(d)
```

Переворачивание строк

```julia
reverse(d,dims=1)
```

Переворачивание столбцов

```julia
reverse(d,dims=2)
```

## Матричное умножение, единичная матрица, скалярное произведение

```julia
A = rand(1:10,(2,3))
```

```julia
B = rand(1:10,(3,4))
```

$\bf{A}$ - матрица $2\times3$, $\bf{B}$ - матрица $3\times4$. Таким образом, их произведение матриц будет $2\times4$ матрицей.

```julia
A*B
```

---
Единичная матрица любого размера может быть получена с помощью функции `Matrix ()` и с помощью специального элемента `I`

```julia
Matrix{Int}(I, 3, 3)
```

```julia
Matrix{Float64}(I, 3, 3)
```

---
Как мы знаем из определения обратной матрицы, что $\bf{A} \bf{B}=\bf{I}$ если $\bf{B}$ есть обратная матрица к $\bf{A}$, т.е. $\bf{B}=\bf{A}^{-1}$. Мы можем проверить, создав случайную матрицу, вычислив ее обратную, а затем умножив исходную матрицу на обратную. Обратите внимание, что мы не получим точную единичную матрицу из-за ограничения числовой точности, но диагональные элементы будут очень близки к 1, а недиагональные элементы будут очень близки к нулю.

```julia
A = rand(1:10,(3,3))
```

```julia
B = inv(A)
```

```julia
A*B
```

---
Скалярное произведение двух векторов вычисляется функцией `dot ()`

```julia
X = [2, 4, -5]
```

```julia
Y = [1,-1,3]
```

```julia
dot(X,Y)
```

```julia
X'*Y
```

```julia
X'Y
```

#### Решение линейных систем

Задача $Ax = b$ для квадратной матрицы $ A $ решается с помощью функции `\`.

```julia
A = rand(3, 3)
```

```julia
x = fill(1.0, 3)
```

```julia
b = A*x
```

```julia
A\b
```

### Упражнения

#### 1 
Умножьте вектор `v` скалярно сам на себя и сохраните результат в `dot_v`.

```julia
v = [1,2,3]
```

```julia

```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 2 
А теперь умножьте `v` матрично на себя (внешнее произведение) присвоив результат переменной `outer_v`

```julia

```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```