# Структуры данных

Как только мы начнем работать со множеством однотипных данных, нам будет удобней хранить данные в структурах, таких как массивы или словари (а не просто создавать множество переменных). 

Типы охватываемых структур данных : 
1. [Кортежи](#Кортежи)
1. [Именованные кортежи](#Именованные-кортежи)
1. [Словари](#Словари)
1. [Массивы](#Массивы)


<br>
Кортежи и массивы - это упорядоченные последовательности элементов (поэтому мы можем индексировать их). Словари и массивы являются изменяемыми. 
Мы объясним это подробнее ниже!

## Кортежи (Tuples)

Мы можем создать кортеж, заключив упорядоченную коллекцию элементов в `( )`.

Синтаксис: 

```julia
(item1, item2, ...)
```

```julia
favoritelang = ("Python","Julia","R")
```

```julia
typeof(favoritelang)
```

```julia
myfavoriteanimals = ("penguins", "cats", "sugargliders")
```

### Indexing
Мы можем индексировать в этот кортеж,

```julia
myfavoriteanimals[1]
```

### Immutability
Кортежи неизменны, мы не можем обновить их

```julia
myfavoriteanimals[1] = "otters"
```

### Iteration
Мы можем перебирать кортеж.

```julia
for lang in favoritelang
    print(lang)
    print(", ")
end
```

### Несколько полезных методов (общих для всех структур данных)
- `isempty` - проверяет, пуста ли структура данных,
- `length` - возвращает длину структуры данных,
- `in` - проверяет принадлежность к элементу,
- `unique` - возвращает коллекцию уникальных элементов,
- `reduce` - сокращает данную коллекцию` itr` с заданным бинарным оператором `op`,
- `maximum` (или` minimum`) - возвращает наибольший (или наименьший) результат вызова функции `fun` для каждого элемента` itr`

```julia
isempty(favoritelang)
```

```julia
length(favoritelang)
```

```julia
"Java" in favoritelang
```

```julia
"Julia" in favoritelang
```

```julia
favoritelang2 = ("Python","Julia","R","Julia","Python");
println("Printing all the elements:")
for lang in favoritelang2
    print(lang);print(", ")
end
println("\n\nNow printing only the unique elements:")
for lang in unique(favoritelang2)
    print(lang);print(", ")
end
```

```julia
reduce(*,favoritelang)
```

```julia
t = (2,3,-4);
reduce(+,t)
```

```julia
# Python is the longest string, so we expect the answer to be 6
maximum(length,favoritelang)
```

```julia
# R is the shortest string, so we expect the answer to be 1
minimum(length,favoritelang)
```

## Именованные кортежи

Как вы могли догадаться, `NamedTuple`ы это такие же `Tuple`ы за исключением того, что каждый элемент дополнительно имеет имя! У них есть специальный синтаксис, использующий `=` внутри кортежа:

```julia
(name1 = item1, name2 = item2, ...)
```

```julia
myfavoriteanimals = (bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Как и обычные `Tuples`, `NamedTuples` упорядочены, так что мы можем получить их элементы с помощью индексации:

```julia
myfavoriteanimals[1]
```

Они также добавляют особую возможность доступа к значениям по их имени:

```julia
myfavoriteanimals.bird
```

## Словари

Если у нас есть наборы данных, связанных друг с другом, мы можем сохранить эти данные в словаре. Мы можем создать словарь, используя функцию `Dict()`.

Syntax:

```julia
Dict(key1 => value1, key2 => value2, ...)
```

Хорошим примером является список контактов, где мы связываем имена с номерами телефонов.

```julia
phonebook = Dict("Jenny" => "867-5309", "Ghostbusters" => "555-2368")
```

В этом примере каждое имя и номер являются парой «ключ» и «значение». Мы можем получить номер Дженни (значение), используя связанный ключ

```julia
phonebook["Jenny"]
```

Мы можем добавить еще одну запись в этот словарь следующим образом

```julia
phonebook["Kramer"] = "555-FILK"
```

Давайте посмотрим, как сейчас выглядит наша телефонная книга...

```julia
phonebook
```

Мы можем удалить Крамера из нашего списка контактов и одновременно получить его номер, используя `pop!`

```julia
pop!(phonebook, "Kramer")
```

```julia
phonebook
```

В отличие от кортежей и массивов, словари не упорядочены. Таким образом, мы не можем индексировать их.

```julia
phonebook[1]
```

В этом примере `julia` думает, что вы пытаетесь получить доступ к значению, связанному с ключом `1`.

### Некоторые полезные методы
- `keys`
- `values`
- `pairs`
- `haskey`

```julia
keys(phonebook)
```

```julia
values(phonebook)
```

```julia
pairs(phonebook)
```

```julia
haskey(phonebook,"Kramer")
```

```julia
haskey(phonebook,"Ghostbusters")
```

### Функция `merge()`
Эта функция создает объединенную коллекцию из заданных коллекций. При необходимости типы результирующей коллекции будут приведены, чтобы соответствовать типам объединенных коллекций. Если такой же ключ присутствует в другой коллекции, значением для этого ключа будет значение, которое он имеет в последней перечисленной коллекции.

```julia
a = Dict("foo" => 0.0, "bar" => 42.0);
b = Dict("baz" => 17, "bar" => 13.0);
```

```julia
a
```

```julia
b
```

```julia
# Note the order of merging
merge(a, b)
```

```julia
# Now the order is reversed
merge(b,a)
```

Если `merge ()` используется с аргументом функции "объединитель", то значения с одним и тем же ключом будут объединены с помощью функции объединения.

```julia
a = Dict("foo" => 0.0, "bar" => 42.0);
b = Dict("baz" => 17, "bar" => 13.0);

# Using a +(sum) function as the combiner
merge(+,a,b)
```

### Массивы

Как и в других программах, под массивами в Julia понимается коллекция упорядоченных элементов, размещённая в многомерной сетке. Векторы и матрицы являются частными случаями массивов.


В Julia предусмотрено достаточно много способов создать массив (с учётом типа, значения, размерности и пр.), но по моему скромному мнению в большинстве случаев удобнее пойти путём Matlab-подобных сред и инициализировать массив через перечисление его элементов в квадратных скобках, разделяя строки знаком " ; ". Чтобы обратиться к элементу, укажите его индекс в квадратных скобках, если массив многомерный, перечислите индексы через запятую. ***Индексация начинается с единицы.*** Для облегчения доступа индекс последнего элемента хранится в переменной end. 

Синтаксис: <br>
```julia
[item1, item2, ...]
```

Создание пустого массива с абстрактным типом

```julia
my_array = []
```

Создание пустого массива с конкретным типом

```julia
my_array = (Integer)[]
```

```julia
my_array = (Float64)[]
```

Массив из трех элементов (здесь тип будет определен автоматически Julia)

```julia
array1 = [1, 2, 3]
```

Это вектор-столбец. Мы также можем создать вектор-строку, опуская запятые.

```julia
array2 = [1 2 3]
```

Если мы смешиваем различные типы элементов в массиве, тогда массивы примут тип элемента, который является *более общим*. В приведенном ниже примере два значения будут иметь тип `Int64`, а последнее - тип `Float64`. Массив вернет тип `Float64`.

```julia
array3 = [1, 2, 3.0]
```

```julia
array4 = ["one",pi,42];
for i in array4
    println("Type of \'$i\' is $(typeof(i))")
end
```

Мы можем создавать массивы, состоящие из более чем одной строки или столбца элементов. Мы можем создавать многомерные массивы. Для этого вставляем квадратные скобки и манипулируем запятыми и точками с запятой.

Обратите внимание на отсутствие запятых между внутренними вложенными квадратными скобками

```julia
array5 = [[1, 2, 3] [4, 5, 6] [7, 8, 9]]
```

Если мы хотим сначала заполнить элементы по строкам, мы используем точку с запятой

```julia
array6 = [[1 2 3]; [4 5 6]; [7 8 9]]
```

Вот некоторые из базовых функций для работы с массивами.
+ **length(A)** - число элементов.
+ **ndims(A)** - число размерностей.
+ **size(A)** - кортеж размерностей.
+ **size(A, n)** - размерность в заданном направлении.
+ **copy(A)** - создание копии массива. 

```julia
array6
```

Проверяем, сколько у нас элементов

```julia
length(array6)
```

Проверка размера (размерности)

```julia
size(array6)
```

Создание трехмерного массива случайных чисел с плавающей запятой

```julia
a = rand(3,2,2)
```

```julia
ndims(a)
```

### Специальные случаи массивов - `ones`, `zeros`, `fill`

```julia
ones(2,3)
```

```julia
zeros(3,3)
```

```julia
fill([1,2],3)
```

### Reshaping, transposing, and slicing/indexing

```julia
a = collect(1:12)
```

```julia
reshape(a,(2,6))
```

```julia
reshape(a,(3,4))
```

```julia
b = reshape(a,(3,4))
```

```julia
c=transpose(b)
```

```julia
b'
```

Обратите внимание на тип транспонированного массива. Он является частью типа `LinearAlgebra`.

Индексирование - 2-я строка, 3-й столбец

```julia
c[2,3]
```

Весь 3-й ряд

```julia
c[3,:]
```

Весь 2-ой столбец

```julia
c[:,2]
```

Блок, состоящий из 2-й и 4-й строк и 1-го и 3-го столбца

```julia
c[[2,4],[1,3]]
```

Символьные значения в массиве. Например, мы можем создать массив ингредиентов для кекса "Нежность"

```julia
ingridients = ["Meal", "Eag", "Sugar", "Butter", "Milk"]
```

`1` в `Array{String,1}` означает что мы задали одномерный вектор. `Array{String,2}` будет 2d матрицей и т.д. `String` это тип входящих элементов.

Еще мы можем сохранить последовательность чисел

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

```julia
mixture = [1, 1, 2, 3, "Пэрри", "Утконос"]
```

Задав массив, мы можем получить отдельные фрагменты данных из этого массива путем индексации в массиве. Например, если мы хотим, получить третий компонент нашего `ingridients`, мы напишем

```julia
ingridients[3]
```

Мы можем использовать индексирование для редактирования существующего элемента массива

```julia
ingridients[3] = "Fish"
```

Julia использует **индексацию начинающуюся с единицы** (как в FORTRAN), а не с нуля как в С или Python.

Мы также можем редактировать массив, используя `push!` и `pop!`. `push!` aдобавляет элемент в конец массива, а `pop!` удаляет последний элемент массива. 

Мы можем добавить еще один номер в нашу последовательность Фибоначчи

```julia
push!(fibonacci, 21)
```

а затем удалить его

```julia
pop!(fibonacci)
```

```julia
fibonacci
```

Массивы могут иметь произвольное количество измерений, а также могут хранить другие массивы.

Например, следующие массивы:

```julia
numbers = [[1, 2, 3], [4, 5], [6, 7, 8, 9]]
```

```julia
favorites = [["koobideh", "chocolate", "eggs"],["penguins", "cats", "sugargliders"]]
```

### Упражнения

#### 3.1 
Создайте массив, `a_ray`, Спомощью следующего кода:

```julia
a_ray = [1, 2, 3]
```

Добавьте число `4` в конец масива, а затем удалите его.

```julia

```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3]
```

#### 3.2 
Попробуйте добавить "Emergency" в `phonebook` задав значение `string(911)`, а потом попробуйте

```julia
phonebook["Emergency"] = 911
```

Что тут не так?

```julia

```

#### 3.3 
Создайте новый словарь `flexible_phonebook` где номер Дженни будет целочисленным, а у Охотников за привидениями - строковым используя следующий код

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

```julia

```

<!-- grade_id: cell-66f1d636c676c2f2, points: 1, grade: true, solution: false -->
```julia
@assert flexible_phonebook == Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

#### 3.4 
Добавьте туда ключ "Emergency" со значением `911` (целочисленное).

```julia

```

<!-- grade_id: cell-ee862ab938680259, points: 1, grade: true, solution: false -->
```julia
@assert haskey(flexible_phonebook, "Emergency")
```

<!-- grade_id: cell-c8c40ababe7cdfce, points: 2, grade: true, solution: false -->
```julia
@assert flexible_phonebook["Emergency"] == 911
```

#### 3.5 
Почему в `flexible_phonebook` целочисленное добавляется, а в `myphonebook` нет? Как нам следует инициализировать `myphonebook`, чтоб он мог принимать целочисленные значения?

```julia

```