# Функции

- Объявление функции
- Утинная типизация в Julia
- Mutating vs. non-mutating functions
- Функции высшего порядка

## Объявление функции

Julia дает нам несколько разных способов написать функцию. Первый требует ключевых слов `function` и` end`

```julia
function sayhi(name)
    println("Hi $name, it's great to see you!")
end
```

```julia
function f(x)
    x^2
end
```

Мы можем вызвать любую из этих функций следующим образом:

```julia
sayhi("C-3PO")
```

```julia
f(42)
```

В качестве альтернативы, мы могли бы объявить любую из этих функций в одной строке

```julia
sayhi2(name) = println("Hi $name, it's great to see you!")
```

```julia
f2(x) = x^2
```

```julia
sayhi2("R2D2")
```

```julia
f2(42)
```

Наконец, мы могли бы объявить их как «анонимные» функции

```julia
sayhi3 = name -> println("Hi $name, it's great to see you!")
```

```julia
f3 = x -> x^2
```

```julia
sayhi3("Chewbacca")
```

```julia
f3(42)
```

### [Утиная типизация](https://ru.wikipedia.org/wiki/%D0%A3%D1%82%D0%B8%D0%BD%D0%B0%D1%8F_%D1%82%D0%B8%D0%BF%D0%B8%D0%B7%D0%B0%D1%86%D0%B8%D1%8F) в Julia

---
*Если нечто выглядит как утка, плавает как утка и крякает как утка, то это, вероятно, и есть утка.*

*If it looks like a duck, swims like a duck and quacks like a duck, then it probably is a duck.*

[Утиный тест](https://ru.wikipedia.org/wiki/%D0%A3%D1%82%D0%B8%D0%BD%D1%8B%D0%B9_%D1%82%D0%B5%D1%81%D1%82)

---

Функции Julia будут работать только с теми входными данными, которые имеют смысл. Например, теперь `sayhi` работает над именем этого второстепенного телевизионного персонажа, записанного как целое число.

```julia
sayhi(55595472)
```

А `f` будет работать и на матрице.

```julia
A = rand(3, 3)
```

```julia
f(A)
```

`f` также будет работать со строкой типа "hi", потому что` * `определено для строковых входов как конкатенация строк.

```julia
f("hi")
```

С другой стороны, `f` не будет работать с вектором. В отличие от `A ^ 2`, который является четко определенным, значение` v ^ 2` для вектора, `v`, не является четко определенной алгебраической операцией.

```julia
v = rand(3)
```

```julia
f(v)
```

## Mutating vs. non-mutating functions

По соглашению, функции, сопровождаемые `!`, изменяют свое содержимое, а функции без `!`, не делают этого. 

Например, давайте посмотрим на разницу между `sort` и` sort! `.

```julia
v = [3, 5, 2]
```

```julia
sort(v)
```

```julia
v
```

`sort (v)` возвращает отсортированный массив, который содержит те же элементы, что и `v`, но` v` остается без изменений.

С другой стороны, когда мы запускаем `sort! (v)`, содержимое v сортируется в массиве `v`.

```julia
sort!(v)
```

```julia
v
```

## [Функции высшего порядка](https://ru.wikipedia.org/wiki/%D0%A4%D1%83%D0%BD%D0%BA%D1%86%D0%B8%D1%8F_%D0%B2%D1%8B%D1%81%D1%88%D0%B5%D0%B3%D0%BE_%D0%BF%D0%BE%D1%80%D1%8F%D0%B4%D0%BA%D0%B0)

---
*Функция высшего порядка -- в программировании функция, принимающая в качестве аргументов другие функции или возвращающая другую функцию в качестве результата. Основная идея состоит в том, что функции имеют тот же статус, что и другие объекты данных.*

---

### map

`map` - это функция высшего порядка в Julia, которая *принимает функцию* в качестве одного из своих входных аргументов. Затем `map` применяет эту функцию к каждому элементу структуры данных, которую вы передаете. Например, выполнение

```julia
map(f, [1, 2, 3])
```
даст вам выходной массив, где функция `f` была применена ко всем элементам `[1, 2, 3]`
```julia
[f(1), f(2), f(3)]
```

```julia
f(x) = x^2
```

```julia
map(f, [1, 2, 3])
```

Здесь мы возвели в квадрат все элементы вектора `[1, 2, 3]`, а не сам вектор. 

Для этого мы могли бы передать `map` анонимную функцию, а не именованную функцию, такую как

```julia
x -> x^3
```

через

```julia
map(x -> x^3, [1, 2, 3])
```

так возводятся в куб все элементы вектора `[1, 2, 3]`

### broadcast

`broadcast` - это еще одна функция высшего порядка. `broadcast` - это обобщение` map`, поэтому он может делать все, что может делать `map`, и даже больше. Синтаксис для вызова `broadcast` такой же, как и для вызова` map`

```julia
broadcast(f, [1, 2, 3])
```

и снова мы применили `f` (возведение в кварат) ко всем элементам `[1, 2, 3]`- на этот раз путем "трансляции" `f`. 

Синтаксическим сахаром для вызова `broadcast` является размещение `.` между именем функции и её входными аргументами. 

Например, 
```julia 
broadcast (f, [1, 2, 3]) 
``` 
совпадает с 
```julia 
f. ([1, 2, 3]) 
```

```julia
f.([1, 2, 3])
```

Еще раз обратите внимание, насколько это отличается от вызова 
```julia 
f ([1, 2, 3]) 
``` 
Мы можем возвести в квадрат каждый элемент вектора, но мы не можем возвести в квадрат вектор!

Чтобы понять суть, давайте посмотрим на разницу между

```julia
f(A)
```
и
```julia
f.(A)
```
для матрицы `A`:

```julia
A = [i + 3*j for j in 0:2, i in 1:3]
```

```julia
f(A)
```

Как и прежде, мы видим, что для матрицы, `A`,
```
f(A) = A^2 = A * A
``` 

С другой стороны,

```julia
B = f.(A)
```

содержит квадраты всех элементов `A`. 
Этот точечный синтаксис для вещания позволяет нам писать относительно сложные составные поэлементные выражения таким образом, чтобы это выглядело естественным / ближе к математической записи. Например, мы можем написать

```julia
A .+ 2 .* f.(A) ./ A
```

вместо

```julia
broadcast(x -> x + 2 * f(x) / x, A)
```

хотя оба способа отработают одинаково.

Чтобы не ставить перед всеми операторами символ `.`, можно применить макрос `@.`

```julia
@. A + 2 * f(A) / A
```

#### Отличие `map ()` и `broadcast ()`

`map()` и `broadcast()` различаются при работе с несколькими коллекциями разных размеров. 
В то время как `broadcast()` будет пытаться привести все объекты к общему измерению, `map()` будет напрямую применять данную функцию поэлементно.

```julia
map(+, 1, [2,2,2])
```

```julia
broadcast(+, 1, [2,2,2])
```

```julia
map(+, [1,1,1], [2,2,2])
```

### Упражнения

---
#### 1 
Напишите функцию `add_one`, которая добавляет 1 к своему входу.

```julia

```

<!-- grade_id: cell-5119b9e9623c1cb7, points: 1, grade: true, solution: false -->
```julia
@assert add_one(1) == 2
```

<!-- grade_id: cell-50f83d27187a2064, points: 1, grade: true, solution: false -->
```julia
@assert add_one(11) == 12
```

---
#### 2 
Используйте `map` или` broadcast`, чтобы увеличить каждый элемент матрицы `A` на` 1` и присвоить его переменной `A1`.

```julia

```

---
#### 3 
Используйте синтаксис броадкаста с точкой для увеличения каждого элемента матрицы `A1` на `1` и сохранения его в переменной `A2`

```julia

```

<!-- grade_id: cell-f3bd5479679a8fe1, points: 0, grade: true, solution: false -->
```julia
@assert A2 == [3 4 5; 6 7 8;9 10 11]
```