# Библиотеки (Пакеты) 

Julia имеет более 2000 зарегистрированных пакетов, что делает их огромной частью экосистемы Julia. Несмотря на это, экосистеме пакетов еще предстоит многое сделать.

Примечательно, что у нас есть вызовы функций первого класса для других языков, обеспечивающие превосходные интерфейсы сторонних функций. Мы можем легко вызвать Python или R, например, с помощью `PyCall` или` Rcall`. Это означает, что вам не нужно ждать, пока экосистема Julia станет полностью зрелой, а переход на Julia не означает, что вы должны отказаться от своего любимого пакета / библиотеки с другого языка! 

Чтобы увидеть все доступные пакеты, ознакомьтесь с 
- <https://julialang.org/packages/>
- <https://juliahub.com/ui/Home>
- <https://juliaobserver.com/>
- <https://github.com/svaksha/Julia.jl>

При первом использовании пакета в данной установке Julia вам необходимо использовать менеджер пакетов, чтобы явно добавить его:

```julia
import Pkg
Pkg.add("Example")
```

Каждый раз, когда вы используете Julia (например, начинаете новый сеанс в REPL или открываете блокнот в первый раз), вы загружаете пакет с ключевым словом `using`

```julia
using Example
```

В исходном коде `Example.jl` на
https://github.com/JuliaLang/Example.jl/blob/master/src/Example.jl 
мы видим следующую объявленную функцию 
``` 
hello ( who :: String) = "Hello, $who"
``` 
Загрузив `Example`, теперь мы можем вызвать `hello`

```julia
hello("it's me. I was wondering if after all these years you'd like to meet.")
```

Теперь давайте поиграем с пакетом цветов

```julia
Pkg.add("Colors")
```

```julia
using Colors
```

Давайте создадим палитру из 100 разных цветов

```julia
palette = distinguishable_colors(100)
```

а затем мы можем создать рандомную матрицу, используя команду `rand`

```julia
rand(palette, 3, 3)
```

### Упражнения

---
#### 1 
Загрузите пакет Primes (source code at https://github.com/JuliaMath/Primes.jl ).

```julia

```

<!-- grade_id: cell-7c4498760f1c26d7, points: 1, grade: true, solution: false -->
```julia
@assert @isdefined Primes
```

---
#### 2 
Убедитесь, что теперь вы можете использовать функцию `primes`, чтобы все все простые числа до 1 000 000 и сохранить их в переменной `primes_list`

```julia

```

<!-- grade_id: cell-ab26afd81c0a11ae, points: 1, grade: true, solution: false -->
```julia
@assert primes_list == primes(1000000)
```